```python
# Initialize Otter
import otter
grader = otter.Notebook("HW3_Gradescope.ipynb")
```

# Homework 3

## Instructions

1. Do not move or rename the included files. If you move or rename files or folders, the local and Gradescope autograders will not work.
1. Refer to files and variables exactly as they appear. Python is case-sensitive, so `df` and `DF` are different variables.
1. After you answer a question, you may run its public tests.
1. After you answer every question:
    1. Restart your kernel
    1. Run up to the last cell
    1. Save your notebook
    1. Run the last cell, which creates the .zip file for Gradescope in the same folder as this notebook
    1. Upload the .zip file to Gradescope
    1. Make sure your local autograder results match your Gradescope autograder results
1. ***The public tests ensure your answers are formatted correctly but do not ensure your answers are correct.***
1. Gradescope runs hidden tests in addition to the public tests but does not reveal public test results until after the due date. Your grade will be "~/100" until after the due date.

## Packages

```python
import numpy as np
import pandas as pd
import datetime as dt
import yfinance as yf
import pandas_datareader.data as web
import matplotlib.pyplot as plt
%matplotlib inline
```

```python
asgnd_seed = 12000
```

```python
np.random.seed(asgnd_seed)
all_tickers = ['AAPL', 'AXP', 'BA', 'BAC', 'C', 'CAT', 'CSCO', 'CVX', 'DIS', 'GS', 'HD', 'IBM', 'INTC', 'JNJ', 'JPM', 'KO',
            'MCD', 'MMM', 'MRK', 'MSFT', 'NKE', 'PFE', 'PG', 'TRV', 'T', 'UNH', 'V', 'VZ', 'WMT', 'XOM']

stock_list = sorted(np.random.choice(all_tickers,7,replace=False))

print(f'These are the tickers of the five stocks assigned to you: {" ".join(stock_list)}. They are stored in a \
variable called stock_list.')
```

## Questions

---

#### 1. Create a DataFrame called **prices** to store the Adjusted Closing Price from Yahoo!Finance for the seven companies you were assigned above. Their tickers are included in a variable called **stock_list**. The sample period is January 1st, 2010 to March 31st, 2023. Each stock's price should be stored in a column labeled 'ticker', where ticker represents that stock's trading ticker.

_Points:_ 10

```python
start_date = '2010-01-01'
end_date = '2023-03-31'
stock_list = ['BA', 'BAC', 'GS', 'JNJ', 'MCD', 'MRK', 'V']
prices = pd.DataFrame()
for ticker in stock_list:
    data = web.DataReader(ticker, 'yahoo', start_date, end_date)
    prices[ticker] = data.iloc[:, 5]

prices.head()
```

```python
grader.check("q01")
```

---

#### 2. Create a DataFrame named **investments** to store the daily value of a $1 investment made on the last trading day in December 2009 in each of the seven stocks. The investment in each stock should be stored in a column labeled 'ticker', where ticker represents that stock's trading ticker.

_Points:_ 10

```python
...
```

```python
grader.check("q02")
```

---

#### 3. Create a new pandas DataFrame named **ans3** to store the two best stock investments over the entire sample period and their total returns. The shape of the DataFrame should be either (2,1) or (1,2).

_Points:_ 5

```python
...
```

```python
grader.check("q03")
```

---

#### 4. Create a new pandas DataFrame named **ans4** to store the two worst stock investments over the entire sample period and their total returns. The shape of the DataFrame should be either (2,1) or (1,2).

_Points:_ 5

```python
...
```

```python
grader.check("q04")
```

---

#### 5. Create a new DataFrame named **investments_mon** to store the end of the month values for each of the seven investment series created in question 2. Again, each column should be labeled 'ticker', where ticker represents that stock's trading ticker.

_Points:_ 10

```python
...
```

```python
grader.check("q05")
```

---

#### 6. Create a new DataFrame named **returns_mon** to store the monthly simple returns for each of the seven stocks. Each column should be labeled 'ticker', where ticker represents that stock's trading ticker.

_Points:_ 10

```python
...
```

```python
grader.check("q06")
```

---

#### 7. Create a DataFrame named **weights** with three columns: 1) TICKER (set as the index); 2) Random_Weights and 3) Equal_Weights. Use the assigned seed as the random seed when generating the random weights.

_Points:_ 10

```python
np.random.seed(asgnd_seed)
...
```

```python
grader.check("q07")
```

---

#### 8. Create a new DataFrame labeled **portfolios** to compute and store the returns of four portfolios: a) two monthly-rebalanced portfolios allocated to the seven stocks according to the random weights, as well as the equal weights (Label the columns 'RW_MR_Ret' and 'EW_MR_Ret', respectively) and b) two buy-and-hold portfolios allocated to the seven stocks according to the random weights, as well as the equal weights (Label the columns 'RW_BH_Ret' and 'EW_BH_Ret', respectively). Drop any observations with missing values.

_Points:_ 10

```python
...
```

```python
grader.check("q08")
```

---

#### 9. Which of the four portfolios was the best performer over the entire sample period? What was its' cumulative (total) return over the entire sample period? Store the two values in two variables called **port9** and **ret9**, respectively.

_Points:_ 5

```python
...
```

```python
grader.check("q09")
```

---

#### 10. Which of the four portfolios was the worst performer over the sample period January 2022 - March 2023? What was its' cumulative (total) return? Store the two values in two variables called **port10** and **ret10**, respectively.

_Points:_ 5

```python
...
```

```python
grader.check("q10")
```

---

<!-- BEGIN QUESTION -->

#### 11. Create a line plot of the four portfolio cumulative return series. Use "figsize", "grid" and "title" arguments in your plot.

_Points:_ 10

```python
...
```

<!-- END QUESTION -->

---

<!-- BEGIN QUESTION -->

#### 12. Create a bar plot of the four portfolios' annualized monthly standard deviations. Use "figsize", "grid" and "title" arguments in your plot.

_Points:_ 10

```python
...
```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q01 (points: 10)
```
>>> prices.shape == (3335,7)
True
```
```
>>> np.allclose(np.sort(prices.sum().values),\
...             np.sort(np.array([540956.44763184,66546.50459862,625187.58621979,330544.01719284,428807.21240997,166449.45791626,341905.05040455])))
True
```


Tests for q02 (points: 10)
```
>>> prices.shape == (3335,7)
True
```
```
>>> investments.sum().values
array([12824.91413411,  5300.24469   ,  4582.25613663,  7550.31941271,
        9974.98807714,  7516.36093401, 17180.52960308])
```
```
>>> np.allclose(np.sort(investments.sum().values),\
...             np.sort(np.array([12824.91413411,5300.24469,4582.25613663,7550.31941271,9974.98807714,7516.36093401,17180.52960308])))
True
```


Tests for q03 (points: 5)
```
>>> pd.DataFrame(ans3).shape==(1,2) or pd.DataFrame(ans3).shape==(2,1)
True
```


Tests for q04 (points: 5)
```
>>> pd.DataFrame(ans4).shape==(1,2) or pd.DataFrame(ans4).shape==(2,1)
True
```


Tests for q05 (points: 10)
```
>>> investments_mon.shape == (160,7)
True
```
```
>>> np.allclose(np.sort(investments_mon.sum().values),\
...             np.sort(np.array([612.44923998,253.20155388,219.02553774,361.37894469,478.94399309,360.61261972,822.70281313])))
True
```


Tests for q06 (points: 10)
```
>>> returns_mon.shape == (160,7)
True
```
```
>>> np.allclose(np.sort(returns_mon.sum().values),\
...             np.sort(np.array([2.40358268,1.53101898,1.43574277,1.41455807,2.03729064,1.78905321,2.71686657])))
True
```


Tests for q07 (points: 10)
```
>>> weights.shape == (7,2)
True
```
```
>>> np.allclose(np.sort(weights.std().values),\
...             np.sort(np.array([7.23747003e-02, 2.99794481e-17])))
True
```


Tests for q08 (points: 10)
```
>>> portfolios.shape == (159, 4)
True
```
```
>>> np.allclose(np.sort(portfolios.sum().values),\
...             np.sort(np.array([1.89426829,1.79750554,1.90401613,1.81311419])))
True
```


Tests for q09 (points: 5)
```
>>> port9 == 'EW_MR_Ret'
True
```
```
>>> np.allclose(ret9, 4.587479379312097)
True
```


Tests for q10 (points: 5)
```
>>> port10 == 'RW_MR_Ret'
True
```
```
>>> np.allclose(ret10, -0.019727480530448815)
True
```
